```python
# Initialize Otter
import otter
grader = otter.Notebook("lab00.ipynb")
```

# Lab 0: Introduction

Welcome to Codebreaking at Cal! We're excited to be teaching you all about codebreaking and cryptography this fall. This lab will serve as an example of what the weekly lab will be like, as well as a quick Python/numpy refresher. 

Since Codebreaking at Cal requires 61A-equivalent experience, we will not be covering programming basics. However, we recognize that some people may be out of practice or unfamiliar with Python. As such, this lab will include some basic Python syntax reminders.

**Question 1.1: Variables**

Python does not require variables to be typed, so you can assign them as such: `variable = value`.

```python
# Comments in Python are like this!
# Set the below variable to True to pass the test.
test = ...
```

```python
grader.check("q1_1")
```

**Question 1.2: Functions, Modulo**

Similarly, functions do not require a specified return value type or input typing. They are written like `def funcName(args):`.

The **modulo** operator in Python is `%`. As a reminder, the modulo operator returns the value of a number mod `n` for `number % n`. For example, `7 % 5 = 2`. 

Hint: How might even and odd numbers vary under modulo 2?

```python
# Complete the function to return True if the given input is even!
def isEven(x):
    ...

```

```python
grader.check("q1_2")
```

**Question 1.3: Loops**

Loops in Python are slightly different than other languages in that they iterate through a list rather than over some loop variable. The function `range(start, end, step)` takes in a start and end index, and increments the value by `step` every iteration. It starts at `start` but does not include `end`.

```python
# Complete the for loop to sum numbers from 0 to 10 (exclusive) stepping by 2 each time.
sum = 0
...
```

```python
grader.check("q1_3")
```

# Question 2: NumPy

NumPy is a very useful Python library that is often used in data-intensive work. We will sometimes use it for our labs as well -- so here's quick overview of the library and a very basic function of it. (We will re-introduce necessary functions as needed later).

`np.arange(start, end, step)` functions the same as `range` except returning a numpy array.
`np.sum(arr)` returns the sum of a numpy array or Python list.

You can add/multiply/etc numpy arrays pairwise if they are of the same dimension!

**Question 2.1: NumPy Initialization**

```python
import numpy as np
```

```python
# Create a NumPy array that contains [0,10)
arr = ...

# Sum the array
arr_sum = ...
```

```python
grader.check("q2_1")
```

**Question 2.2: NumPy Arrays**

```python
# Add two arrays together 
arr1, arr2 = np.arange(0, 3), np.arange(10, 13)

output = ...
```

```python
grader.check("q2_2")
```

# Question 3: Syllabus Quiz

There are just a few basic questions we want you to answer to confirm you're familiar with the course policies. You can revisit them over at https://codebreakingatcal.org/about.

**Question 3.1:** What is the lecture time for this course? Please put the number of your answer in the variable below.

1: Monday 5PM - 7PM

2: Tuesday 6:30PM - 8PM

3: Tuesday 5PM - 7PM

4: Friday 11AM - 12PM

```python
q31_answer = ...
```

```python
grader.check("q3_1")
```

**Question 3.2**: How many lectures and labs do you need to attend for full credit in each category? Place a number into each variable below.

```python
required_lectures = ...
required_labs = ...
```

```python
grader.check("q3_2")
```

**Question 3.3**: Which resources are available for students to access? Put all numbers that apply into a list.

1. Lecture 
2. Office Hours 
3. Course Notes
4. Class Discord

```python
# resources = [7, 8, 9] (like this)
resources = ...
```

```python
grader.check("q3_3")
```

# Optional: Examples!

This part of the lab is optional, but serves as a quick glimpse into what you will be building next week! Play around with encryption and decryption using a basic polyalphabetic cipher!

```python
%%capture
# Install our custom library
import sys
!{sys.executable} -m pip install git+https://github.com/CodebreakingAtCal/codebreaking_at_cal.git

import codebreaking_at_cal as cbc
```

```python
# How to use the CBC library:
# cbc.vignere_encrypt(plaintext, key) -> ciphertext
# cbc.vignere_decrypt(ciphertext, key) -> plaintext

message = "YOUR MESSAGE HERE"
key = "YOUR KEY HERE"

ciphertext = cbc.vignere_encrypt(message, key)
ciphertext
```

```python
cbc.vignere_decrypt(ciphertext, key)
```

You're done with Lab 0! This was a short lab, and it's likely many of you felt it was a little underwhelming. Don't worry, the purpose of this lab is just to get used to Jupyter and Python a bit. Go ahead and run the following code to generate a ZIP file to submit to Gradescope.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you have generated the zip file, go to the Gradescope page for this assignment to submit.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: 10)
```
>>> assert test
```


Tests for q1_2 (points: 10)
```
>>> assert isEven(2)
>>> assert not isEven(3)
```


Tests for q1_3 (points: 10)
```
>>> assert sum == (0 + 2 + 4 + 6 + 8)
```


Tests for q2_1 (points: 10)
```
>>> np.testing.assert_array_equal(arr, np.arange(0, 10))
>>> assert arr_sum == np.sum(arr)
```


Tests for q2_2 (points: 10)
```
>>> np.testing.assert_array_equal(output, np.array([10, 12, 14]))
```


Tests for q3_1 (points: 10)
```
>>> assert q31_answer == 2
```


Tests for q3_2 (points: 10)
```
>>> assert required_lectures == required_labs == 10
```


Tests for q3_3 (points: 10)
```
>>> assert resources == [1,2,3,4]
```
