```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

```python
%%capture
import sys
!{sys.executable} -m pip install pycryptodome
```

```python
import math
import matplotlib.pyplot as plt
import numpy as np
import string
import itertools
import re
from functools import reduce
import hashlib
from Crypto.Hash import SHA256
```

```python
%%capture
# Install our custom library
import sys
!{sys.executable} -m pip install git+https://github.com/CodebreakingAtCal/codebreaking_at_cal.git

import codebreaking_at_cal
```

# Lab 2: Information Theory
Contributions From: Ryan Cottone, Imran Khaliq-Baporia

Welcome to Lab 2! In this lab, we will demonstrate the basics of information theory, and some techniques used to break Enigma.

As a quick review from lecture and Note 2, the *entropy* of a probability distribution is derived from the following formula:

$$ H(x) = -\sum_{i=0} x_i \log_2(x_i) $$

where $x_i$ represents the probability of event $i$ occuring. It is used to quantify how uncertain a random variable is.

**Question 1.1**: Implement calculateEntropy, which takes in a array of probabilities and returns its entropy.

*HINT: math.log(x, base) returns $\log_{base}(x)$*

```python
def calculateEntropy(probabilityDistribution):
    total = 0
        
    for probability in probabilityDistribution:
        if probability == 0:
            continue
        ...
    
    return -total
    
```

```python
grader.check("q1_1")
```

It is a lot easier to compute the entropy of a uniform distribution (an array of equal probabilities). We can reduce our equation to the following for a uniform distribution with $k$ entries and $x$ probability:

$$H(X_{uniform}) = -k \cdot x\log_2(x)$$

```python
def calculateUniformEntropy(probability, k):
    return -k*probability*math.log(probability, 2)
```

Let's visualize what entropy actually looks like. Below is a graph with the x-axis representing the probability of tossing heads in an unfair coin. You'll notice that the entropy peaks at 0.5, or when the coin is fair. This is at the point where the outcome is most uncertain. Contrast this to when there's a 90% chance to hit heads -- we can be pretty certain about the outcome (lower entropy).

```python
plt.plot(np.arange(0.1,.9001,0.001), [calculateEntropy([x, 1-x]) for x in np.arange(0.1,.9001,0.001)])
plt.show()
```

Specifically, we have one bit of entropy for a single fair coinflip. This corresponds to a sample space of size 2 with uniform probability. We can do no better but to guess the result each time.

In the context of cryptographic keys, we'd need to guess approximately $2^k$ keys for a system with $k$ bits of entropy.

Let's demonstrate the entropy of some ciphers we tried before:

```python
sample_text = ""
with open('sample.txt', 'r') as file:
    sample_text = file.read().replace('\n', '').replace(" ", "")
    
sample_text = codebreaking_at_cal.clean_text(sample_text)
sample_text = re.sub('[^a-z]*', '', sample_text.lower())
```

```python
# Takes in two NumPy arrays and returns their TVD.
def analyzeFrequency(text):
    diff = abs(freq1 - freq2)
    return sum(diff)/2
```

```python
# Takes in a Caesar ciphertext and returns the map of shift -> probability
def shiftProbabilityCaesar(ciphertext):
    ciphertext = re.sub('[^a-z]*', '', ciphertext.lower())
    
    arr = np.array([])
        
    for i in range(0,26):
        arr = np.append(arr, 2**(1/codebreaking_at_cal.analyze_frequency(codebreaking_at_cal.caesar_decrypt(ciphertext, i))))
    
    arr = arr/sum(arr)
    return arr
```

```python
caesarEntropy = calculateEntropy(shiftProbabilityCaesar(codebreaking_at_cal.caesar_encrypt(sample_text, 1)))

k = 20
baseEntropy = calculateUniformEntropy(1/(26**k), 26**k)

print('A Caesar encrypted message of length', len(sample_text), "has", caesarEntropy, "bits of entropy.")
print('A random message of length', k, "has", baseEntropy, "bits of entropy.")
```

Read through those functions and try to understand what they do. In the last cell, we calculated the entropy of a Caesar cipher on a very long text and the entropy of a completely random alphabetic string of length 20 (far, far less than the Caesar text).

To put these numbers into perspective (note that e+28 = $\cdot$ 10^28):

```python
print('There are ', 2**caesarEntropy, ' items in the "sample space" for a caesar encrypted message of length', len(sample_text))
print('There are ', 2**baseEntropy, ' items in the "sample space" for a random message of length', k)
```

Take a look at the graph of the entropy of a Caesar cipher on a $k$ length text (represented by the x-axis):

```python
plt.plot(np.arange(1, 175),
         [calculateEntropy(shiftProbabilityCaesar(codebreaking_at_cal.caesar_encrypt(sample_text[:x], 9))) for x in np.arange(1,175
                                                                                                     )])
plt.show()
```

Note the entropy actually decreasing as the text size gets larger! Can you think of why this is the case?

Now for the graph of a uniformly random string:

```python
y = 10
plt.plot(range(1, y),
         [calculateUniformEntropy(1/26,26**x) for x in range(1,y)])
plt.xticks(range(1,y))
plt.yscale('log', base=10)
plt.show()
```

Note the exponentially increasing y-axis!

We can see that vigenere is better but not by much.

```python
def calculateVigenereEntropy(text, keylen):
    texts = ['']*keylen
    for i in range(len(text)):
        texts[i%keylen] += text[i]
    probabilities = [shiftProbabilityCaesar(x) for x in texts]
        
    return sum([calculateEntropy(x) for x in probabilities])

key = "ab"    
calculateVigenereEntropy(codebreaking_at_cal.vignere_encrypt(sample_text[:50], key), len(key))
```

**Question 1.2**: Is the entropy of a keylength 1 vigenere cipher equal to a Caesar cipher? Enter True or False.

```python
concept_check = ...
```

```python
grader.check("q1_2")
```

The following function generates a random key for use in later functions.

```python
def genRandomKey(length):
    return ''.join(np.random.choice([x for x in string.ascii_lowercase], length))

```

What happens if we use a Vigenere cipher with a random key? Check out the graph of key length versus entropy below:

*NOTE: The orange bar represents the entropy of a 102-long string*

```python
cap = 102
plt.plot(range(1,cap,10),
         [calculateVigenereEntropy(codebreaking_at_cal.vignere_encrypt(sample_text[:cap], genRandomKey(x)), x) for x in range(1,cap,10)])
plt.plot(range(1, cap,10),
         [calculateUniformEntropy(1/26**cap,26**cap) for x in range(1,cap,10)])
plt.show()
```

Do you see why an increasingly long keylength corresponds to more entropy? What do you think happens as the keylength approaches the message size?

# One-Time Pad

Let's introduce a quite surprising idea -- there exists a cipher that is **provably unbreakable** and not any more complicated than the Vigenere cipher!

We will define provably unbreaking, also known as perfectly secret, to be a ciphertext with entropy equal to a random bitstring of the same length. This means that the ciphertext yields no more information than a string of truly random ones and zeroes. You have zero recourse but to try and brute force every possibility (we will see that even this is useless).

Enter the One-Time Pad. As it's name suggests, we can only use it a single time before it becomes insecure again. A one-time pad operates by pairing each letter/bit of the plaintext with a truly random key letter/bit and shifting/XOR-ing. Basically, if we used Vigenere cipher with a truly random key of the same length as the plaintext, we would have a perfectly secure cipher.

You may be wondering why cryptography is even an active field if such a perfect cipher exists. Well, there are considerable drawbacks to such a cipher, namely the onerous requirement of a very large amount of perfectly random key material.

**Question 2.1**: Which of the following message/key pairs are secure?

1: abcd, efgh

2: abcd, b

3: abcd, longkey

Answer in terms of a list, e.g. [1,2,3] means all three answers are correct.

```python
correct = ...
```

```python
grader.check("q2_1")
```

```python
def otp_encrypt(text, key):
    if len(text) != len(key):
        raise Exception()

    return codebreaking_at_cal.vignere_encrypt(text, key)
    
```

```python
def otp_decrypt(text, key):
    if len(text) != len(key):
        raise Exception()
        
    return codebreaking_at_cal.vignere_decrypt(text, key)
    
```

Let's graph the entropy of sample_text encoded with a OTP. We know that each 

```python
def count_letters(text):
    counts = {}
    text = text.lower()
    
    for letter in string.ascii_lowercase:
        counts[letter] = 0
    
    for letter in text:
        if (letter in string.ascii_lowercase): 
            counts[letter] += 1
    
    return counts

def calculate_proportions(text): # Coded for you
    counts = count_letters(text).values()
    nparr = np.fromiter(counts, dtype=float)
    return nparr / sum(counts)

plt.bar([string.ascii_lowercase[i] for i in range(26)], calculate_proportions(otp_encrypt(sample_text, genRandomKey(len(sample_text)))))
plt.show()
```

Looks pretty uniform, right? Assuming a random shift at each step, we have 26 possibilities at each index. For a string of length $k$, that means $26^k$ possible keys. Moreover, we will see that the properties of a OTP make it impossible to decrypt at all, not just hard.

The crux of have a separate shift for every letter is that you can decrypt it into **every possible string** by changing the key. With traditional ciphers like Vigenere or Caesar, there were a bounded number of different strings possible with a key of a certain length. This is because eventually the key repeats with the same plaintext, and you can't map these two occurences to different outcomes. With a OTP, that is very much possible.

Take for example the ciphertext "AUVDGP". 

Say we found our possible key "ABCDEF": decrypt("AUVDGP", "ABCDEF") == "ATTACK"

Instead, say we found a possible key "XQQZTM": decrypt("AUVDGP", "XQQZTM") == "DEFEND"

Which is correct? There's no way to tell.

Let's revisit why this is the case. The formula for encryption can be thought of the following, where adding is pairwise by letter:

PLAINTEXT + KEY = CIPHERTEXT

To find the key for any arbitrary plaintext from the ciphertext, simply solve the following:

KEY = CIPHERTEXT - PLAINTEXT = vignere_decrypt(CIPHERTEXT, PLAINTEXT)

## Random Number Generation

It's time to apply what we've learned about information theory to the study of random number generation! As you might recall from lecture, a **PRNG** (pseudo-random number generator) is a computer program that *deterministically* outputs "random looking" numbers. A *cryptographically secure* PRNG fulfills the randomness requires of an PRNG as well as some security considerations (rollback resistance, resistant to internal state leak, etc).

```python
def H(x):
    if type(x) != bytes:
        if type(x) == str:
            x = bytes(x, 'utf-8')
        elif type(x) == int:
            x = x.to_bytes(len(bin(x))-2, byteorder='big')
    
    h = SHA256.new()
    h.update(x)
    return h.digest().hex()
```

```python
def hashPRNG(seed):
    state = seed
                
    def generate(bits):
        nonlocal state
        
        state = int(H(state),16)  
        return state & ((1 << bits) - 1)
        
    return generate
```

Let's try to analyze and break a weakly-seeded PRNG. The above CSPRNG is (relatively) secure, so we'll focus on the initial seeding.

Alice seeds the RNG by flipping a coin for each bit (H = 1, T = 0) for 32 bits. This would be decently secure if Mallory hadn't tampered with the coin beforehand, making it 99% chance for heads and 1% for tails!

```python
calculateEntropy([0.99, 0.01])*32
```

We see that the entropy of this seeding process is barely 2.5 bits, not the 32 bits Alice expects!

```python
def weightedCoin():
    if np.random.rand() > 0.99:
        return 0
    else: 
        return 1
```

Since 1s are far more common than 0s, we expect our number to be much larger than smaller. So if we start from MAX_NUMBER and decremenet, we should hit our seed quickly! This is roughly analogous to guessing 1 for every bit, then guessing more 0s slowly, which is what we expect to be fastest.

In order to test our guess, we compare an output of 100 bits from the PRNG to the output of 100 bits from a fresh PRNG seeded with our guess. If they equal, there is a very very high change the underlying seed is equal ($1 - 2^{-100}$ probability).

```python
seed = int(''.join([str(weightedCoin()) for _ in range(32)]), 2)
gen = hashPRNG(seed)
hundred_bit_output = gen(100)

for i in range(2**32, -1, -1):
    testPRNG = hashPRNG(i)
    hundred_bits = testPRNG(100)
    
    if hundred_bit_output == hundred_bits:
        print('Found at', i)
        break
```

Guessing 32 bits would take a considerable amount of time on a Jupyter notebook! You've (probably) just cracked the PRNG in a few seconds instead.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you have generated the zip file, go to the Gradescope page for this assignment to submit.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: 10)
```
>>> assert calculateEntropy([0.5, 0.5]) == 1
>>> assert calculateEntropy([1, 0]) == 0
```


Tests for q1_2 (points: 10)
```
>>> assert concept_check == True
```


Tests for q2_1 (points: 10)
```
>>> assert correct == [1, 3]
```
